# Core Language Part 1

For the first part of the homework, you are asked to write out some lines of code and define variables with particular names. In the second half of the homework, you are asked to fill code in functions.

### 1. Create and use two variables (1.5 points)

**1a.** (0.5 point) Create a variable called `var1` that equals the integer 10. Create another variable called `var2` that equals the float 2.5. Multiply `var1` by `var2` and call the result `var3`. 

*Don't forget:* You will need to delete the line with `raise NotImplementedError`, write your own solution, and then re-run the cell before running the test cell. Each time you change code in a cell, you  will need to re-run that cell before running any other cells that depend on it!

<!-- grade_id: cell-6a3e3f6b30b2ee0e, points: 0, grade: false, solution: true -->
```python
# Input the code part of your answer here.
var1 = 10
var2 = 2.5

var3 = var1 * var2
```

<!-- grade_id: cell-6dff7ca5e48eb93c, points: 0.5, grade: true, solution: false -->
```python
# Test the value of each variable
assert(var1 == 10)
assert(var2 == 2.5)
assert(var3 == 25)

print('Success!')
```

Output:
```
Success!
```

**1b.** (0.5 point) Write (using Markdown) what the type of `var3` should be and why.

<!-- grade_id: cell-341bf7ec6f32d4fb, points: 0.5, grade: true, solution: true -->
var3 is a float type since it's multiplying an integer with a float, the result will output as a float to try maintain precision in case the value has deicimal points and is therefore not an integer.

**1c.** (0.5 point) Check if `var3` is greater than 15 and save result to the variable name `checkvar3`.

<!-- grade_id: cell-ae656e231c94fee7, points: 0, grade: false, solution: true -->
```python
checkvar3 = var3 > 15
```

<!-- grade_id: cell-2f400f395a98fbd5, points: 0.5, grade: true, solution: false -->
```python
assert(checkvar3)

print('Success!')
```

Output:
```
Success!
```

### 2. Operations on strings (1 point)

**2a.** (0.5 point) Define two strings: `string1` contains the word "ink" and `string2` contains the word "jet". Combine the two words together and call the resultant variable `printer`.

<!-- grade_id: cell-2ea540ee892203eb, points: 0, grade: false, solution: true -->
```python
string1 = 'ink'
string2 = 'jet'
printer = string1 + string2
```

<!-- grade_id: cell-f06f27643c8964e6, points: 0.5, grade: true, solution: false -->
```python
# test type of variables
assert(isinstance(string1, str))
assert(isinstance(string2, str))
assert(isinstance(printer, str))

# test final result
assert(printer == 'inkjet')

print('Success!')
```

Output:
```
Success!
```

**2b.** (0.5 point) Split the string by the letter "k" and save the resultant list in a variable named `reslist`.

<!-- grade_id: cell-4f3584f327c70d2f, points: 0, grade: false, solution: true -->
```python
reslist = printer.split('k')

```

<!-- grade_id: cell-8d90e109198c0c97, points: 0.5, grade: true, solution: false -->
```python
assert(isinstance(reslist, list))
assert(reslist[0] == 'in')
assert(reslist[1] == 'jet')

print('Success!')
```

Output:
```
Success!
```

### 3. Make a list (1.5 points)

**3a.** (0.5 point) Create a variable called `alist` that contains, in order, the string 'dog', the boolean True, and the number 34. Write your answer in the following cell.

<!-- grade_id: cell-b8940cdbf0e33284, points: 0, grade: false, solution: true -->
```python
alist = ['dog', True, 34]
```

<!-- grade_id: cell-0b9aa531121ee3a4, points: 0.5, grade: true, solution: false -->
```python
# This line tests to see if the variable you created, called `alist`, is of type list, and contains the correct element types
assert(isinstance(alist, list))
assert(isinstance(alist[0], str))
assert(isinstance(alist[1], bool))
assert(isinstance(alist[2], int))

# This line tests to make sure the correct elements are contained in the list
assert(alist[0] == 'dog')
assert(alist[1] == True)
assert(alist[2] == 34)

print('Success!')
```

Output:
```
Success!
```

**3b.** (0.5 point) Change the last entry of the list to equal the list `[2,5,9]`.

<!-- grade_id: cell-d21e7dc5ecf4f34f, points: 0, grade: false, solution: true -->
```python
alist[2] = [2, 5, 9]

```

<!-- grade_id: cell-efd6296b634211fd, points: 0.5, grade: true, solution: false -->
```python
assert(isinstance(alist[2], list))
assert(alist[2] == [2,5,9])

print('Success!')
```

Output:
```
Success!
```

**3c.** (0.5 point) Check if index 1 of `alist` is equivalent to the string "True" and save the result as variable `test`.

<!-- grade_id: cell-6126d685e7ef4b2c, points: 0, grade: false, solution: true -->
```python
test = alist[1] == 'True'

```

<!-- grade_id: cell-5956761dfba66c52, points: 0.5, grade: true, solution: false -->
```python
assert(not test)

print('Success!')
```

Output:
```
Success!
```

### 4. Make a dictionary (1 point)

Make a dictionary that contains the capital cities of a given state. Instead of all 50 states, just include Washington, Texas, and California. The dictionary must be called `state_capitals` and should be able to used as follows:
    
    state_capitals['Texas'] = 'Austin'


<!-- grade_id: dictionary, points: 0, grade: false, solution: true -->
```python
# Used this website as guidance for creating this dictionary
# https://docs.python.org/3/tutorial/datastructures.html#dictionaries

state_capitals = {'Texas':'Austin', 'Washington':'Olympia', 'California':'Sacramento'}
```

<!-- grade_id: test_dictionary, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous code. This cell should NOT give any errors when it is run."""

# makes sure your code returns a dictionary
assert(isinstance(state_capitals, dict))

# checks that the capitals can be accessed correctly
assert(state_capitals['Texas'] == 'Austin')
assert(state_capitals['Washington'] == 'Olympia')
assert(state_capitals['California'] == 'Sacramento')

# if nothing throws an assertion, good job!
print("Success!")
```

Output:
```
Success!
```

**References:**
* Python-3 Startup Guide
* https://docs.python.org/3/tutorial/datastructures.html#dictionaries

```python

```