### Core Language 3 Part 1

### 1. Compute the length of a given list or string (2 points)

Define a function that computes the length of a given list or string. (It is true that Python has the len() function built in, but do not use it for your solution — use a loop instead.)

For example:

    >>> length(['string1', 'string2'])
    2
    >>> length('string1')
    7

<!-- grade_id: cell-7cc34eb8c688c3a5, points: 0, grade: false, solution: true -->
```python
def length(inputs):
    """Compute the length of a given list or string, but without using `len`.
    
    Parameters
    ----------
    inputs : string or list
        
        
    Returns
    -------
    integer that is either the number of characters in the input string or the number of 
    entries in the input list.
    """
    ### BEGIN SOLUTION
    
    length_count = 0
    for string in inputs:
        length_count += 1
    
    return length_count
    
print(length('wednesday'))
print(length(['monday', 'tuesday', 'wednesday', 'thursday', 'friday', 'saturday', 'sunday']))
```

Output:
```
9
7
```

<!-- grade_id: cell-3870df85ff34e2bb, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
from unittest import TestCase
assert_equal = TestCase().assertEqual

# Check a list
inputs = ['apple', 'orange']
assert_equal(length(inputs), 2)

# Check a string
inputs = 'tomato'
assert_equal(length(inputs), 6)

# What about a string with a space in it?
inputs = 'bacon and eggs'
assert_equal(length(inputs), 14)

# Also will look for not using the built-in len() function
len_old = len
len = None
inputs = 'tomato'
assert_equal(length(inputs), 6)
len = len_old

print("Success!")
```

Output:
```
Success!
```

### 2a-b. Convert list of words into list of integers of length of words (2 points)

Write a series of functions that maps a list of words into a list of integers representing the lengths of the corresponding words. Write them in two different ways: 

1. using a for-loop
1. using list comprehensions. 

Each function will return a list of word lengths, for example,

    >>> method1(['list', 'of', 'words'])
    [4, 2, 5]

<!-- grade_id: forloop, points: 0, grade: false, solution: true -->
```python
# part a
def method1(input1):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using a for loop.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    ### BEGIN SOLUTION
    
    list_lengths = []
    for list_item in input1:
        length_of_item = len(list_item)
        list_lengths += [length_of_item]
    return list_lengths

print(method1(['monday', 'tuesday', 'wednesday', 'thursday', 'friday', 'saturday', 'sunday']))
```

Output:
```
[6, 7, 9, 8, 6, 8, 6]
```

<!-- grade_id: test_forloop, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method1(words), [3, 3, 5])
# Also will check that a for loop was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: listcomprehension, points: 0, grade: false, solution: true -->
```python
# part b
def method2(input2):
    """Returns a list of integers representing the lengths of the corresponding input list of words.
    Do this using a list comprehension.
    
    Parameters
    ----------
    input1 : list of str
        inputs to be counted   
        
    Returns
    -------
    list of integers counting the elements in the input list
    """
    ### BEGIN SOLUTION

    lists = [len(list_item) for list_item in input2]
    return lists

print(method2(['monday', 'tuesday', 'wednesday', 'thursday', 'friday', 'saturday', 'sunday']))
```

Output:
```
[6, 7, 9, 8, 6, 8, 6]
```

<!-- grade_id: test_listcomprehension, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
# check with a list
words = ['one', 'two', 'three']
assert_equal(method2(words), [3, 3, 5])
# Also will check that a list comprehension was used
print('success!')
```

Output:
```
success!
```

<!-- grade_id: test_all, points: 1, grade: true, solution: false -->
```python
"""(1 point) Test code for the previous function. This cell should NOT give any errors when it is run."""
words = ['five', 'eight', 'infinity']
# Check that the result is the same for all methods
assert_equal(method1(words), method2(words))
print('success!')
```

Output:
```
success!
```

### 3. Number of integers to reach 1 million (2 points)

Calculate how many integers, starting from 1 and counting each consecutive integer, are needed to add up to more than 1 million. Use a `while` function to do this. You shouldn't need to subtract 1 from your final answer to get this to work out.

<!-- grade_id: cell-3e995c9f072905fb, points: 0, grade: false, solution: true -->
```python
def ints_to_million():
    '''Find the number of consecutive integers (starting from 1) that are needed to add 
    together to sum to more than 1 million.'''
    
    ### BEGIN SOLUTION
    
    integer = 0
    total = 0

    while total <= 1000000:
        integer += 1
        total += integer   
    return integer
    
print(ints_to_million()) 
```

Output:
```
1414
```

<!-- grade_id: cell-c17f04140621a1f3, points: 2, grade: true, solution: false -->
```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""

# save the number of integers
numints = ints_to_million()

assert(numints*(numints+1)/2 == 1000405)

print('success!')
```

Output:
```
success!
```

### 4. Encode and decode using ROT-13 (2 points)

In cryptography, a Caesar cipher is a very simple encryption techniques in which each letter in the plain text is replaced by a letter some fixed number of positions down the alphabet. For example, with a shift of 3, A would be replaced by D, B would become E, and so on. The method is named after Julius Caesar, who used it to communicate with his generals. ROT-13 ("rotate by 13 places") is a widely used example of a Caesar cipher where the shift is 13. In Python, the key for ROT-13 may be represented by means of the following dictionary:

    key = {'a':'n', 'b':'o', 'c':'p', 'd':'q', 'e':'r', 'f':'s', 'g':'t', 'h':'u', 
           'i':'v', 'j':'w', 'k':'x', 'l':'y', 'm':'z', 'n':'a', 'o':'b', 'p':'c', 
           'q':'d', 'r':'e', 's':'f', 't':'g', 'u':'h', 'v':'i', 'w':'j', 'x':'k',
           'y':'l', 'z':'m', 'A':'N', 'B':'O', 'C':'P', 'D':'Q', 'E':'R', 'F':'S', 
           'G':'T', 'H':'U', 'I':'V', 'J':'W', 'K':'X', 'L':'Y', 'M':'Z', 'N':'A', 
           'O':'B', 'P':'C', 'Q':'D', 'R':'E', 'S':'F', 'T':'G', 'U':'H', 'V':'I', 
           'W':'J', 'X':'K', 'Y':'L', 'Z':'M'}

Your task in this exercise is to implement an encoder/decoder of ROT-13. Once you're done, you will be able to read the following secret message:

       Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!

Note that since English has 26 characters, your ROT-13 program will be able to both encode and decode texts written in English. You should be able to do this using the properties of a dictionary.

```python
key = {'a':'n', 'b':'o', 'c':'p', 'd':'q', 'e':'r', 'f':'s', 'g':'t', 'h':'u', 
       'i':'v', 'j':'w', 'k':'x', 'l':'y', 'm':'z', 'n':'a', 'o':'b', 'p':'c', 
       'q':'d', 'r':'e', 's':'f', 't':'g', 'u':'h', 'v':'i', 'w':'j', 'x':'k',
       'y':'l', 'z':'m', 'A':'N', 'B':'O', 'C':'P', 'D':'Q', 'E':'R', 'F':'S', 
       'G':'T', 'H':'U', 'I':'V', 'J':'W', 'K':'X', 'L':'Y', 'M':'Z', 'N':'A', 
       'O':'B', 'P':'C', 'Q':'D', 'R':'E', 'S':'F', 'T':'G', 'U':'H', 'V':'I', 
       'W':'J', 'X':'K', 'Y':'L', 'Z':'M'}

def rot13(letter):
    """Swap letter indicated in key.
    
    Parameters
    ----------
    letter : string of length 1
        letter to be swapped using the key dictionary
        
    Returns
    -------
    letter
    """
    # References: https://docs.python.org/3/library/functions.html#len, https://www.w3schools.com/python/ref_dictionary_get.asp,
    # https://docs.python.org/3/library/stdtypes.html#dict.get
    
    phrase_decoded = '' # empty string
    
    for item in letter:
        phrase_decoded += key.get(item, item) # the second 'item' is necessary for spaces in the phrase since spaces do not exist in the key
                                               # without second item it would return an error    
    return phrase_decoded
    
print(rot13('Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!'))
```

Output:
```
Caesar cipher? I much prefer Caesar salad!
```

```python
"""(2 points) Test code for the previous function. This cell should NOT give any errors when it is run."""
phrase = 'Pnrfne pvcure? V zhpu cersre Pnrfne fnynq!'

translatedchars = list(map(rot13, phrase))  # map connects a function with 
                                            # each letter in phrase.
translated = ''.join(translatedchars)  # we got a bunch of characters back. Join combines them
                                       # into a string

# recode
retranslatedchars = list(map(rot13, translated))
retranslated = ''.join(retranslatedchars)
    
# Check decode
import unittest
test = unittest.TestCase()
test.assertEqual(translated, "Caesar cipher? I much prefer Caesar salad!")
test.assertEqual(retranslated, phrase)  # check that it works both ways

print("Success!")
```

Output:
```
Success!
```

**References:**
* https://docs.python.org/3/tutorial/controlflow.html#for-statements
* https://docs.python.org/3/tutorial/datastructures.html#list-comprehensions
* https://docs.python.org/3/tutorial/controlflow.html#defining-functions
* https://docs.python.org/3/library/stdtypes.html#dict.get
* Python-3 Startup Guide
* M3 Coding With Python
* M2 Coding With Python
* https://www.w3schools.com/python/python_lists_comprehension.asp
* https://www.w3schools.com/python/ref_dictionary_get.asp
* https://docs.python.org/3/library/functions.html#len

```python

```